### Adaboost

This activity focuses on using the `AdaBoostClassifier` and the performance resulting from changing the base classifier that is used.  As discussed in the lectures, adaptive boosting is a successive reweighting of data using a set number of estimators.  These weighted estimators are what form the ensemble, and the predictions are a result of a weighted combination of the estimators.  

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)
- [Problem 4](#-Problem-4)
- [Problem 5](#-Problem-5)

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.ensemble import AdaBoostClassifier
from sklearn.tree import DecisionTreeClassifier
from sklearn.linear_model import LogisticRegression
from sklearn.preprocessing import StandardScaler
from sklearn.pipeline import Pipeline
from sklearn.model_selection import GridSearchCV, train_test_split
import warnings
warnings.filterwarnings('ignore')
```

```python
df = pd.read_csv('data/fetal.zip', compression = 'zip')
```

```python
X = df.drop('fetal_health', axis = 1).values
y = df['fetal_health']
```

```python
X_train, X_test, y_train, y_test = train_test_split(X, y, random_state=42)
```

[Back to top](#-Index)

### Problem 1

#### `AdaBoostClassifier`

**10 Points**

What is the default estimator in the `AdaBoostClassifier`?  Instantiate the default estimator with the correct hyperparameters to `ans1` below.

<!-- grade_id: cell-610a6433eeb011fa, points: 0, grade: false, solution: true -->
```python
### GRADED
ans1 = ''
    
### BEGIN SOLUTION
ans1 = DecisionTreeClassifier(max_depth=1)
### END SOLUTION

### ANSWER CHECK
ans1
```

Output:
```
DecisionTreeClassifier(max_depth=1)
```

<!-- grade_id: cell-254a544f694d9ceb, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans1_ = DecisionTreeClassifier(max_depth=1)
#
#
#
assert ans1.max_depth == ans1_.max_depth
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### Fitting the Ensemble

**10 Points**

Below, use the `AdaBoostClassifier` to fit the data.  Use all default settings and and assign the accuracy of the model on the test data to `model_1_acc` below.

<!-- grade_id: cell-92ed88043191a639, points: 0, grade: false, solution: true -->
```python
### GRADED
model_1 = ''
model_1_acc = ''
    
### BEGIN SOLUTION
model_1 = AdaBoostClassifier().fit(X_train, y_train)
model_1_acc = model_1.score(X_test, y_test)
### END SOLUTION

### ANSWER CHECK
print(model_1_acc)
```

Output:
```
0.881578947368421
```

<!-- grade_id: cell-a4e2adfd28e5c9fe, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
model_1_ = AdaBoostClassifier().fit(X_train, y_train)
model_1_acc_ = model_1_.score(X_test, y_test)
#
#
#
assert model_1_acc_ == model_1_acc
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 3

#### Grid Searching the Ensemble

**10 Points**

As the documentation states [here](https://scikit-learn.org/stable/modules/ensemble.html#usage), the main parameters to search are the number of estimators and the complexity of the base estimator.  Create a parameter grid that considers the following parameters:

- *number of estimators*: 100, 200
- *max_depths*: 1, 2, 3

as `params` below.  Use this with the `AdaBoostClassifier` to grid search named `tree_grid` on the train data.  Assign the score on the test data as `grid_acc`.  Be sure to set the `random_state = 42` in your `AdaBoostClassifier`.

**NOTE:** This is computaitonally expensive. It may take up to two minutes for the answer check (print(grid_acc) ) results to appear. It is also advised that you do NOT run the blank cell underneath (the grading cell) until you see the results from your answer.


<!-- grade_id: cell-0510fdbbf90c3779, points: 0, grade: false, solution: true -->
```python
### GRADED
params = ''
tree_grid = ''
grid_acc = ''

### BEGIN SOLUTION
params = {'n_estimators': [100, 200], 'estimator__max_depth': [1, 2, 3]}
tree_grid = GridSearchCV(AdaBoostClassifier(estimator=DecisionTreeClassifier(), random_state = 42), param_grid=params).fit(X_train, y_train)
grid_acc = tree_grid.score(X_test, y_test)
### END SOLUTION

### ANSWER CHECK
print(grid_acc)
```

Output:
```
0.9417293233082706
```

<!-- grade_id: cell-2c2d4e716fcfe852, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
params_ = {'n_estimators': [100, 200], 'estimator__max_depth': [1, 2, 3]}
tree_grid_ = GridSearchCV(AdaBoostClassifier(estimator=DecisionTreeClassifier(), random_state=42), param_grid=params_).fit(X_train, y_train)
grid_acc_ = tree_grid_.score(X_test, y_test)
#
#
#
assert grid_acc == grid_acc_
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 4

#### A Different Base Estimator

**10 Points**

Consider using a different base estimator such as `LogisticRegression` estimator.  Explore the neighbors parameters with 

- `C = [.001, 0.01, 0.1, 1.0, 10.0]`

Create a `Pipeline` that scales the data first and then implements an `AdaBoostClassifier` with `random_state = 42` and a Logistic Regression model.  Grid search the pipeline with a grid and assign the score on the test data to `score2`. 

**Note:** Again, this one is computationally expensive. Be patient with the results, and do NOT run the blank cell underneath (the grading cell) until you see the results from your answer.


<!-- grade_id: cell-2bbf690915a58f45, points: 0, grade: false, solution: true -->
```python
### GRADED
score2 = ''

    
### BEGIN SOLUTION
params = {'mod__base_estimator__C': [.001, 0.01, 0.1, 1.0, 10.0]}
p = Pipeline([('scale', StandardScaler()), ('mod', AdaBoostClassifier(base_estimator = LogisticRegression(), random_state = 42)) ])
g = GridSearchCV(p, param_grid=params)
g.fit(X_train, y_train)
score2 = g.score(X_test, y_test)
### END SOLUTION

### ANSWER CHECK
print(score2)
```

Output:
```
0.9078947368421053
```

<!-- grade_id: cell-84b60d9b9e7df6eb, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
params_ = {'mod__base_estimator__C': [.001, 0.01, 0.1, 1.0, 10.0]}
p_ = Pipeline([('scale', StandardScaler()),
             ('mod', AdaBoostClassifier(base_estimator = LogisticRegression(), 
                                       random_state = 42))
             ])
g_ = GridSearchCV(p_,
                param_grid=params_)
g_.fit(X_train, y_train)
score2_ = g_.score(X_test, y_test)
#
#
#
assert score2 == score2_
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 5

#### Evaluating the models

**10 Points**

Which model performed the best on the test data?

- `a`: Base `AdaBoostClassifier`
- `b`: Grid Searched Tree Model
- `c`: Grid Searched Logistic Model
- `d`: None of the above

Assign your answer as a string to `ans5` below.

<!-- grade_id: cell-12e99140e358ffcc, points: 0, grade: false, solution: true -->
```python
### GRADED
ans5 = ''
    
### BEGIN SOLUTION
ans5 = 'b'
### END SOLUTION

### ANSWER CHECK
print(ans5)
```

Output:
```
b
```

<!-- grade_id: cell-7ee93976c0a78160, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans5_ = 'b'
#
#
#
#
assert ans5 == ans5_
### END HIDDEN TESTS
```