## Codio Activity 7.5: Calculating Multiple Loss Functions

**Expected Time = 60 minutes**

**Total Points = 20**

A third loss function mentioned is the Huber loss function.  This is notable for its resistance to extreme values and is defined as a piecewise function:


$${\displaystyle L_{\delta }(y,f(x))={\begin{cases}{\frac {1}{2}}(y-f(x))^{2}&{\textrm {for}}|y-f(x)|\leq \delta ,\\\delta \,(|y-f(x)|-{\frac {1}{2}}\delta ),&{\textrm {otherwise.}}\end{cases}}}$$

In this activity, you will compute and compare the results of minimizing the mean squared error, mean absolute error, and huber loss functions.  

NOTE: If the formula is not rendering correctly (overlapping text), double-click in this cell and then Shift-Enter to reload the cell.


## Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)

```python
import seaborn as sns
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import plotly.express as px
import plotly.graph_objects as go
from scipy.optimize import minimize
```

### The tips data

For this exercise, the tips dataset from the lectures will be used, and you are to predict the tip amount given the total bill.  

```python
tips = sns.load_dataset('tips')
```

```python
tips.head()
```

Output:
```
total_bill   tip     sex smoker  day    time  size
0       16.99  1.01  Female     No  Sun  Dinner     2
1       10.34  1.66    Male     No  Sun  Dinner     3
2       21.01  3.50    Male     No  Sun  Dinner     3
3       23.68  3.31    Male     No  Sun  Dinner     2
4       24.59  3.61  Female     No  Sun  Dinner     4
```

```python
X = tips['total_bill']
y = tips['tip']
```

[Back to top](#Index:) 

## Problem 1

### Create a Huber Loss function

**10 Points**

Using the formula repeated below for the Huber loss, complete the function that returns the sum of the Huber Loss equation.

$${\displaystyle L_{\delta }(y,f(x))={\begin{cases}{\frac {1}{2}}(y-f(x))^{2}&{\textrm {for}}|y-f(x)|\leq \delta ,\\\delta \,(|y-f(x)|-{\frac {1}{2}}\delta ),&{\textrm {otherwise.}}\end{cases}}}$$

Your function will take in both a value for $\theta$ and for $\delta$.  Set the value of $\delta = 1.5$ by default.

<!-- grade_id: cell-9a9b87589083f36b, points: 0, grade: false, solution: true -->
```python
#GRADED

def huber_loss(theta, delta):
    """
    This function accepts a value for theta
    and returns the sum of the huber loss.
    
    Arguments
    ---------
    theta: float
           Values to use for parameter
           of regression model.
           
    delta: float
           Value for delta in Huber Loss
            
    Returns
    -------
    huber: np.float
         Sum of huber loss
    """

### BEGIN SOLUTION
def huber_loss(theta, delta = 1.5):
    y_pred = theta*X
    y_err = np.abs(y - y_pred) 
    return sum(np.where(y_err <= delta, 1/2*(y_err)**2, delta*(y_err - 1/2*delta)))
### END SOLUTION

huber_loss(8)
```

Output:
```
56561.369999999995
```

<!-- grade_id: cell-5921ecbcb23b0fb0, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def huber_loss_(theta, delta = 1.5):
    y_pred = theta*X
    y_err = np.abs(y - y_pred) 
    return sum(np.where(y_err <= delta, 1/2*(y_err)**2, delta*(y_err - 1/2*delta)))
ans_ = huber_loss_(8)
stu = huber_loss(8)
#
#
#
assert type(ans_) == type(stu), 'Make sure you return a float'
assert ans_ == stu, 'Your function does not return the same solution for theta = 8'
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 2

### Minimizing Huber Loss

**10 Points**

Use the `minimize` function imported from `scipy.optimize` to determine the optimal value for $\theta$ based on the Huber Loss with $\delta = 1.5$.  Assign your results as type `np.float` to `theta_huber` below with `x0 = .5`.  

<!-- grade_id: cell-91b58b91bc727853, points: 0, grade: false, solution: true -->
```python
### GRADED

theta_huber = ''

### BEGIN SOLUTION
minimum_theta = minimize(huber_loss, x0 = .5)
theta_huber = minimum_theta.x[0]
### END SOLUTION

# Answer check
print(type(theta_huber))
print(theta_huber)
```

Output:
```
<class 'numpy.float64'>
0.14626752601211537
```

<!-- grade_id: cell-837b15e800192fb7, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def huber_loss_(theta, delta = 1.5):
    y_pred = theta*X
    y_err = np.abs(y - y_pred) 
    return sum(np.where(y_err <= delta, 1/2*(y_err)**2, delta*(y_err - 1/2*delta)))

minimum_theta_ = minimize(huber_loss_, x0 = .5)
theta_huber_ = minimum_theta_.x[0]

#
#
#
assert type(theta_huber_) == type(theta_huber)
assert minimum_theta.success == minimum_theta_.success
assert np.isclose(theta_huber_, theta_huber, atol=1e-6)
np.testing.assert_equal(minimum_theta_, minimum_theta)

### END HIDDEN TESTS
```

```python

```