## Codio Activity 1.1: Python Coding

**Expected time: 30 minutes**

**Total Points: 12**
    
    This assignment focuses on creating DataFrames using a dictionary, url, and local file.
    
## Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)

```python
import pandas as pd
```

[Back to top](#Index:) 

## Problem 1

### Creating a DataFrame

**4 Points**

Use a dictionary to create a DataFrame. Assign your results to the variable total_widgets below.


<!-- grade_id: cell-def2f9b9bb27ac1d, points: 0, grade: false, solution: true -->
```python
###GRADED
#Create a dataframe from the the widgets dictionary

#Create the dictionary of simulated data
data = {
    'WidgetA': [10, 9, 7, 8], 
    'WidgetB': [11, 12, 30, 1]
}

#And then pass it to the pandas DataFrame constructor total_widgets

total_widgets = None

### BEGIN SOLUTION

total_widgets = pd.DataFrame(data)

### END SOLUTION

#Printing the total widgets 
print(total_widgets)
```

Output:
```
WidgetA  WidgetB
0       10       11
1        9       12
2        7       30
3        8        1
```

<!-- grade_id: cell-a151493f2710ac6f, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from pandas.testing import assert_frame_equal

data = {
    'WidgetA': [10, 9, 7, 8], 
    'WidgetB': [11, 12, 30, 1]
}

total_widgets1 = pd.DataFrame(data)

assert_frame_equal(total_widgets, total_widgets1)

### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 2

### Create a dataframe from url

**4 Points**

Now, use a url to load data to a DataFrame. The url below links to a .csv file in a github repository containing data about diamonds. Assign your resulting DataFrame to the variable diamonds below.
    

```python
url = 'https://raw.githubusercontent.com/mwaskom/seaborn-data/master/diamonds.csv'
```

<!-- grade_id: cell-8bb615f9c4e19312, points: 0, grade: false, solution: true -->
```python
### GRADED

diamonds = ''

### BEGIN SOLUTION
diamonds = pd.read_csv(url)
### END SOLUTION

# Answer check
print(type(diamonds))
diamonds.head()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
   carat      cut color clarity  depth  table  price     x     y     z
0   0.23    Ideal     E     SI2   61.5   55.0    326  3.95  3.98  2.43
1   0.21  Premium     E     SI1   59.8   61.0    326  3.89  3.84  2.31
2   0.23     Good     E     VS1   56.9   65.0    327  4.05  4.07  2.31
3   0.29  Premium     I     VS2   62.4   58.0    334  4.20  4.23  2.63
4   0.31     Good     J     SI2   63.3   58.0    335  4.34  4.35  2.75
```

<!-- grade_id: cell-32a08456b4bdff4a, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
diamonds_ = pd.read_csv(url)
#
#
#
assert type(diamonds_) == type(diamonds)
pd.testing.assert_frame_equal(diamonds_, diamonds)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

##  Problem 3

### Create a dataframe by importing data from an file.

**4 Points**

Now, use the datafile located in the data folder named nba.csv. Assign your results as a DataFrame to the variable nba below.

<!-- grade_id: cell-7d6b9e515bd3b13f, points: 0, grade: false, solution: true -->
```python
### GRADED

nba = ''

### BEGIN SOLUTION
nba = pd.read_csv('data/nba.csv')
### END SOLUTION

# Answer check
print(type(nba))
nba.head()

```

Output:
```
<class 'pandas.core.frame.DataFrame'>
            Name            Team  Number Position   Age Height  Weight  \
0  Avery Bradley  Boston Celtics     0.0       PG  25.0    6-2   180.0   
1    Jae Crowder  Boston Celtics    99.0       SF  25.0    6-6   235.0   
2   John Holland  Boston Celtics    30.0       SG  27.0    6-5   205.0   
3    R.J. Hunter  Boston Celtics    28.0       SG  22.0    6-5   185.0   
4  Jonas Jerebko  Boston Celtics     8.0       PF  29.0   6-10   231.0   

             College     Salary  
0              Texas  7730337.0  
1          Marquette  6796117.0  
2  Boston University        NaN  
3      Georgia State  1148640.0  
4                NaN  5000000.0
```

<!-- grade_id: cell-a0cb61ae1f5a3eb1, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
nba_ = pd.read_csv('data/nba.csv')
#
#
#
assert type(nba_) == type(nba)
pd.testing.assert_frame_equal(nba, nba_)
### END HIDDEN TESTS

```