### Codio Activity 22.8: Hyperparameter Tuning with `keras`

**Expected Time = 60 minutes**

**Total Points = 40**

This activity focuses on using hyperparameter tuning with the keras library.  There are two ways to perform a grid search with keras, and you will implement both. While `keras_tuner` was discussed in the lectures, here you will use the scikit-learn wrapper for keras where you can grid search the parameters using `GridSearchCV`.  You will implement this with the `KerasClassifier` to build some basic models on the wine dataset.  

#### Index

- [Problem 1](#-Problem-1)
- [Problem 2](#-Problem-2)
- [Problem 3](#-Problem-3)
- [Problem 4](#-Problem-4)

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
import tensorflow as tf
import warnings
warnings.filterwarnings('ignore')

#from tensorflow.keras.wrappers.scikit_learn import KerasClassifier
from scikeras.wrappers import KerasClassifier
from tensorflow.keras.layers import Dense
from tensorflow.keras.models import Sequential
from tensorflow.keras.utils import to_categorical
from sklearn.datasets import load_wine
from sklearn.preprocessing import StandardScaler
from sklearn.model_selection import train_test_split, GridSearchCV
```

### The Data

Below, the wine dataset is loaded, split, and scaled.  

```python
wine = load_wine(as_frame=True)
```

```python
wine.frame.head()
```

Output:
```
alcohol  malic_acid   ash  alcalinity_of_ash  magnesium  total_phenols  \
0    14.23        1.71  2.43               15.6      127.0           2.80   
1    13.20        1.78  2.14               11.2      100.0           2.65   
2    13.16        2.36  2.67               18.6      101.0           2.80   
3    14.37        1.95  2.50               16.8      113.0           3.85   
4    13.24        2.59  2.87               21.0      118.0           2.80   

   flavanoids  nonflavanoid_phenols  proanthocyanins  color_intensity   hue  \
0        3.06                  0.28             2.29             5.64  1.04   
1        2.76                  0.26             1.28             4.38  1.05   
2        3.24                  0.30             2.81             5.68  1.03   
3        3.49                  0.24             2.18             7.80  0.86   
4        2.69                  0.39             1.82             4.32  1.04   

   od280/od315_of_diluted_wines  proline  target  
0             
```

```python
X = wine.data
y = to_categorical(wine.target)
```

```python
X_scaled = StandardScaler().fit_transform(X)
```

```python
X_scaled.shape
```

Output:
```
(178, 13)
```

[Back to top](#-Index)

### Problem 1

#### The Build Function

**10 Points**

To use the `KerasClassifier` you first need to write a function that creates a keras model and takes in arguments for the parameters you wish to search. Below is a very basic example.  Note that rather than 

```python
def build_function(neurons = 1):
    model = Sequential()
    model.add(Dense(neurons, activation = 'relu'))
    model.add(Dense(3, activation = 'softmax'))
    model.compile(...)
    return model
```

Complete the `build_function` below with appropriate arguments for `model.compile()`.

After you are done, call the function `build_function` with argument `neurons = 100` and assign it to the variable `ans`.



<!-- grade_id: cell-a2bce77d17112fd1, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
def build_function(neurons = 1):
    model = ''
    #model.compile()
    return model

ans = None

    
### BEGIN SOLUTION
tf.random.set_seed(42)
def build_function(neurons = 1):
    model = Sequential()
    model.add(Dense(neurons, activation = 'relu', input_dim=13))
    model.add(Dense(neurons, activation = 'relu'))
    model.add(Dense(3, activation = 'softmax'))
    model.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
    return model
ans = build_function(neurons = 100)
### END SOLUTION

### ANSWER CHECK
build_function
```

Output:
```
<function __main__.build_function(neurons=1)>
```

<!-- grade_id: cell-817b7e106ad4c353, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
def build_function_(neurons = 1):
    model = Sequential()
    model.add(Dense(neurons, activation = 'relu', input_dim=13))
    model.add(Dense(neurons, activation = 'relu'))
    model.add(Dense(3, activation = 'softmax'))
    model.compile(loss = 'categorical_crossentropy', metrics = ['accuracy'])
    return model
ans = build_function_(neurons = 100)
stu = build_function(neurons = 100)
#
#
#
assert ans.layers[0].units == stu.layers[0].units
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 2

#### Creating the `KerasClassifier`

**10 Points**

Now, use the `build_function` to instantiate the `KerasClassifier` as `model` with `epochs = 10`. Fit the model on the data `X_scaled` and `y`, and assign the accuracy as `score_1` below.  Note that this will use 1 node as defaulted in the `build_function`.

<!-- grade_id: cell-0275116e4f507015, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
model = ''
score_1 = ''

    
### BEGIN SOLUTION
tf.random.set_seed(42)
model = KerasClassifier(build_fn=build_function, 
                        epochs = 10, 
                       verbose = 0)
history = model.fit(X_scaled, y)
model.score(X_scaled, y)
### END SOLUTION

### ANSWER CHECK
print(model.score(X_scaled, y))
```

Output:
```
0.398876404494382
```

<!-- grade_id: cell-7696f2d368bb0be4, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
model_ = KerasClassifier(build_fn=build_function_, 
                        epochs = 10, 
                       verbose = 0)
history = model_.fit(X_scaled, y)
ans = model_.score(X_scaled, y)
stu = model.score(X_scaled, y)
#
#
#
assert ans == stu
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[35], line 12
      8 stu = model.score(X_scaled, y)
      9 #
     10 #
     11 #
---> 12 assert ans == stu
     13 ### END HIDDEN TESTS

AssertionError:
```

[Back to top](#-Index)

### Problem 3

#### Performing the Grid Search

**10 Points**

Now, to perform a grid search you just need to create a dictionary with the hyperparameter named `neurons = [10, 50, 100]`.  Use `GridSearchCV` with the model and appropriate hyperparameter dictionary to fit the data.  

```python
ann_grid = GridSearchCV(model, param_grid = params)
```

<!-- grade_id: cell-8c00e8a1dee417d8, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)
params = ''
ann_grid = ''

    
### BEGIN SOLUTION
tf.random.set_seed(42)
params = {'neurons': [10, 50, 100]}
ann_grid = GridSearchCV(model, param_grid=params)
### END SOLUTION

### ANSWER CHECK
ann_grid
```

Output:
```
GridSearchCV(estimator=KerasClassifier(build_fn=<function build_function at 0x00000286CD523600>, epochs=10, verbose=0),
             param_grid={'neurons': [10, 50, 100]})
```

<!-- grade_id: cell-addbffa1795cc12f, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
params_ = {'neurons': [10, 50, 100]}
ann_grid_ = GridSearchCV(model_, param_grid=params_)
#
#
#
assert ann_grid.get_params()['param_grid'] == ann_grid_.get_params()['param_grid']
### END HIDDEN TESTS
```

[Back to top](#-Index)

### Problem 4

#### Fit and Evaluate the model

**10 Points**

Now, fit and evaluate the model.  Assign your score to `grid_score` below. 

<!-- grade_id: cell-f8e6edb517c71335, points: 0, grade: false, solution: true -->
```python
### GRADED
tf.random.set_seed(42)

grid_score = ''

    
### BEGIN SOLUTION
tf.random.set_seed(42)
params = {'neurons': [10, 50, 100]}
ann_grid = GridSearchCV(model, param_grid=params)
grid_score = ann_grid.fit(X_scaled, y).score(X_scaled, y)
### END SOLUTION

### ANSWER CHECK
print(grid_score)
```

<!-- grade_id: cell-fbf8f8e25550261d, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
tf.random.set_seed(42)
params_ = {'neurons': [10, 50, 100]}
ann_grid_ = GridSearchCV(model_, param_grid=params_)
grid_score_ = ann_grid_.fit(X_scaled, y).score(X_scaled, y)
#
#
#
assert grid_score == grid_score_
### END HIDDEN TESTS
```

Because of the grading enviornment, a more exhaustive search over additional parameters is not an option.  To extend the work here should be straightforward enough, and this is a nice solution to grid searching the hyperparameters of a model. 